// variables6.rs
//
// Execute `rustlings hint variables6` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

const NUMBER: i32 = 3;
fn main() {
    println!("Number {}", NUMBER);
}
// in this exercise, i learnt about constants and that one must always annotate the type.
// Here’s an example of a constant declaration:

// const THREE_HOURS_IN_SECONDS: u32 = 60 * 60 * 3;
// also learnt about Rust’s naming convention for constants 
// is to use all uppercase with underscores between words.
// onto the next🔥
